# Project 1 Notebook

The overview of Project 1 is written here: http://courses.las.illinois.edu/spring2020/stat107/projects/mosaic/.  This notebook is used for coding the project with some provided starter code. :)

## Part 1: Average Image Color

Follow the guide for Project 1 to help complete your `findAverageImageColor` function:

<!-- grade_id: cell-0dbb8ddddc2b1de5, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import skimage
import PIL
import numpy as np

def findAverageImageColor(lab, width, height):
    data = []
    data_f = []
    for x in range(width):
        for y in range(height):
            L, a, b = lab[x][y]
            d = {'L':L,'a':a,'b':b}
            data.append(d)
            
    df = pd.DataFrame(data)
    L_C = df['L'].sum()
    a_C = df['a'].sum() 
    b_C = df['b'].sum()
    
    pixels = width * height
    
    
    L_avg = L_C /pixels
    a_avg = a_C/pixels
    b_avg = b_C/pixels
    
    f = {'L':L_avg,'a':a_avg, 'b':b_avg}
    #data_f.append(f)
    #df_f = pd.DataFrame(data_f)
    return f
            


            
    
    
    
```

### Part 1 -- Test Cases

These test cases help you debug your code to make sure you're finding the correct average values for known images.  After you pass these test cases, then you'll start using your own images! :)

<!-- grade_id: cell-59256d97f4255439, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1c ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

info = '\N{INFORMATION SOURCE}'
%run project1-settings.ipynb

def test_findAverageImageColor(file):
    from PIL import Image
    image = Image.open(file)
    w, h = image.size
    lab = imageToLabArray(image, resize=False)
    return findAverageImageColor(lab, w, h)

test = test_findAverageImageColor('test.png')

assert( test != None ), "Your findAverageImageColor function must return a value (right now it's returning nothing)."
assert( type(test) == type({}) ), f"Your findAverageImageColor function must return a dictionary (right now it's returning {type(test)})."

assert( 'L' in test ), "Your findAverageImageColor must return a value for 'L'."
assert( 'a' in test ), "Your findAverageImageColor must return a value for 'a'."
assert( 'b' in test ), "Your findAverageImageColor must return a value for 'b'."

print(f"{info} Your test.png values: (L={test['L']}, a={test['a']}, b={test['b']})")
assert( abs(test['L'] - 54.244) > 0.001 ), "Your 'L' value is the value of only orange pixels.  Are you sure you are visiting every pixel?"
assert( abs(test['a'] - 59.314) > 0.001 ), "Your 'a' value is the value of only orange pixels.  Are you sure you are visiting every pixel?"
assert( abs(test['b'] - 52.9799) > 0.001 ), "Your 'b' value is the value of only orange pixels.  Are you sure you are visiting every pixel?"

assert( abs(test['L'] - 47.197) < 0.001 ), "Your 'L' value is not correct on test.png."
assert( abs(test['a'] - 49.034) < 0.001 ), "Your 'a' value is not correct on test.png."
assert( abs(test['b'] - 38.609) < 0.001 ), "Your 'b' value is not correct on test.png."


test2 = test_findAverageImageColor('test2.png')

print(f"{info} Your test2.png values: (L={test2['L']}, a={test2['a']}, b={test2['b']})")
assert( abs(test2['L'] - 54.244) < 0.001 ), "Your 'L' value is not correct on test2.png."
assert( abs(test2['a'] - 59.314) < 0.001 ), "Your 'a' value is not correct on test2.png."
assert( abs(test2['b'] - 52.9799) < 0.001 ), "Your 'b' value is not correct on test2.png."


test3 = test_findAverageImageColor('test3.png')

print(f"{info} Your test3.png values: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
assert( abs(test3['L'] - 46.414) < 0.001 ), "Your 'L' value is not correct on test3.png."
assert( abs(test3['a'] - 47.892) < 0.001 ), "Your 'a' value is not correct on test3.png."
assert( abs(test3['b'] - 37.012) < 0.001 ), "Your 'b' value is not correct on test3.png."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print()
print(f"{tada} All tests passed! {tada}")
```

Output:
```
ℹ Your test.png values: (L=47.19722525581814, a=49.03421116311882, b=38.60877549417686)
ℹ Your test2.png values: (L=54.244093289693964, a=59.3141053878179, b=52.979879933089656)
ℹ Your test3.png values: (L=46.41423991872082, a=47.89200069370781, b=37.01198611207544)

🎉 All tests passed! 🎉
```

## Part 1 Image Analysis -- Your Images

Once your function works and passes all the test cases, this code loads all files in the folder specified in `project1-settings.ipynb` and stores the average color value of the image in `tile-images.csv`.

- This code is already complete, but depends on a correct `findAverageImageColor` function!
- Make sure your code is complete before running this function!
- If you have a File Not Found Error, double check the location of your tile images in `project1-settings.ipynb`!


```python
import pandas as pd

import PIL
from PIL import Image

from os import listdir
from os.path import isfile, join
import sys


# Run the `project1-settings.ipynb` notebook:
%run project1-settings.ipynb

def isImageFile(file):
    if file.endswith(".jpg") or file.endswith(".jpeg") or file.endswith(".png") or file.endswith(".gif") or \
       file.endswith(".JPG") or file.endswith(".JPEG") or file.endswith(".PNG") or file.endswith(".GIF"):
        return True
    else:
        return False


# Use `data` to store data about every image:
data = []
    
# Loop through every file:
error = False
files = [tileImageDir + "/" + f for f in listdir(tileImageDir + "/") if isfile(join(tileImageDir + "/", f)) and isImageFile(join(tileImageDir + "/", f))]
print(f"Starting to process {len(files)} files...")
for file in files:
    try:
        # Reduce file to outputSize in pixels
        lab = fileToLabArray(file)
        w = len(lab)
        h = len(lab[0])
                
        # Run `findAverageImageColor` and store the result:
        result = findAverageImageColor(lab, w, h)
        d = { 'file': file, 'L': result['L'], 'a': result['a'], 'b': result['b'] }
        data.append(d)
        
        # Print out message for the user:
        pct = len(data) / len(files) * 100
        sys.stdout.write(f"\r  ... {len(data)} / {len(files)} total files processed ({pct:.2f}%)")
    except ValueError as e:
        print()
        print()
        print(f"ERROR: `{file}` failed to process.")
        print("...try to remove the image and run again.")
        print(repr(e))
        error = True
        break
                
# Save the data as a DataFrame and CSV file:
if not error:
    df = pd.DataFrame(data)
    df.to_csv('tile-images.csv')
    df

    tada = "\N{PARTY POPPER}"
    print("")
    print("")
    print(f"{tada} Image data saved as ""tile-images.csv""")
    print(df)
```

Output:
```
Starting to process 1805 files...
  ... 1805 / 1805 total files processed (100.00%)

🎉 Image data saved as tile-images.csv
                                                   file          L          a  \
0     images//12543290_743782945758728_2050232560_n.png  43.423874   1.603944   
1     images//929215_352643904900366_812094193_n-9-1...  54.047168   8.768161   
2     images//17267774_227185591020106_4151870126842...  53.456438   5.778413   
3      images//11282781_107601682906321_284282954_n.png  57.428358   6.698477   
4     images//929215_352643904900366_812094193_n-9-3...  43.733876   3.814923   
...                                                 ...        ...        ...   
1800   images//1515111_471984069660729_1275938831_n.png  46.017757   1.672017   
1801  images//929215_352643904900366_812094193_n-9-4...  57.341433  17.884570   
1802  images//929215_352643904900366_812094193_n-9-1...  55.214231  -2.963837   
1803  images//929215_352643904900366_812094193_n-9-4...  53.740130 
```

## Turn In and Continue to Part 2

- Make sure to turn in Part 1 by the due date!
- Part 2 will use the CSV file you saved, `tile-images.csv`, to build your image mosaic!

```python

```