# Your Grades and GPA at Illinois!

In this lab, you will explore the GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

# Part 0: Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = "Sogol"
ca1_name = "Katy"
ca2_name = "Emily"


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = "Shawn"
partner1_netid = "coffell2"
partner1_major = "Info Sci"
partner1_qotd_answer = "Shapeshifting"

partner2_name = "michelle"
partner2_netid = "mj15"
partner2_major = "CS + Stats"
partner2_qotd_answer = "Mind reading"

partner3_name = "Finn"
partner3_netid = "FJc33"
partner3_major = "Econ"
partner3_qotd_answer = "Create wealth"

```

## Part 0.5 | Loading the Dataset

For the first part of our lab, we'll be using the GPA dataset we saw in class.

You can download the dataset by following to following link to the github storing the dataset: https://raw.githubusercontent.com/wadefagen/datasets/master/gpa/uiuc-gpa-dataset.csv

To download this file, you will want to **right click** and click **Save As**.

<!-- grade_id: cell-51fef9428e349da7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = pd.read_csv('GPA.csv')
```

## Part 1 | Finding Total GPAs

Add two new columns, `Total Students` and `Average GPA`, that compute the total number of students in each course and the average weighted GPA for each course.

You should not include students with a `W` in either column and feel free to add other columns to help you out!

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python
df["Total Students"] = df["A+"]+df["A"]+df["A-"]+df["B+"]+df["B"]+df["B-"]+df["C+"]+df["C"]+df["C-"]+df["D+"]+df["D"]+df["D-"]+df["F"]

df["Average GPA"] = (df["A+"]*4.0+df["A"]*4.0+df["A-"]*3.67+df["B+"]*3.33+df["B"]*3+df["B-"]*2.67+df["C+"]*2.33+df["C"]*2+df["C-"]*1.67+df["D+"]*1.33+df["D"]*1+df["D-"]*0.67)/ df["Total Students"]


```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(df[
  (df["YearTerm"] == "2010-su") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 410) \
]["Total Students"].sum() == 31), "It appears your \"Total Students\" column is incorrect."

assert(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "CS") & \
  (df["Number"] == 225) \
]["Total Students"].sum() == 697), "It appears your \"Total Students\" column is incorrect."

assert( abs(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 100) \
]["Average GPA"].mean() - 3.247) < 0.01), "It appears your \"Average GPA\" column is incorrect."
```

## Part 2 | Finding the easiest and hardest courses

### Puzzle 2.1: Hardest Courses in the Dataset

Using functions you already know, find the **30 hardest courses** at Illinois based on their "Average GPA".

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hardestCourses = df.nsmallest(30, "Average GPA")
df_hardestCourses


```

Output:
```
Year    Term YearTerm Subject  Number                   Course Title  \
4111   2018    Fall  2018-fa     LER     120    Contemporary Labor Problems   
16583  2016  Spring  2016-sp     LER     100  Introduction to Labor Studies   
24396  2014    Fall  2014-fa     LER     120    Contemporary Labor Problems   
26664  2014  Spring  2014-sp     LER     120    Contemporary Labor Problems   
29461  2013    Fall  2013-fa    MATH       2           Introductory Algebra   
26739  2014  Spring  2014-sp    MATH     220                       Calculus   
1407   2019  Spring  2019-sp     LER     100  Introduction to Labor Studies   
14311  2016    Fall  2016-fa     LER     120    Contemporary Labor Problems   
14579  2016    Fall  2016-fa     MCB     245     Human Anat & Physiol Lab I   
24443  2014    Fall  2014-fa     LLS     296  Bilingual US Latina/o Culture   
24607  2014    Fall  2014-fa     MCB     245     Human Anat & Physiol Lab I   
11512  2017  Spring  2017-sp     LER     100  Introduction 
```

### Puzzle 2.2: Hardest Courses in 2019

Find the **30 hardest courses in 2019** at Illinois based on their "Average GPA".

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_2019 = df[df["Year"]==2019]
df_hardestCourses_2019 = df_2019.nsmallest(30, "Average GPA")
df_hardestCourses_2019


```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
1407  2019  Spring  2019-sp     LER     100   Introduction to Labor Studies   
2244  2019  Summer  2019-su    CHEM     102             General Chemistry I   
1410  2019  Spring  2019-sp     LER     120     Contemporary Labor Problems   
1487  2019  Spring  2019-sp    MATH     112                         Algebra   
2376  2019  Summer  2019-su    SPAN     141        Intro to Spanish Grammar   
918   2019  Spring  2019-sp    ECON     302      Inter Microeconomic Theory   
526   2019  Spring  2019-sp    CHEM     232  Elementary Organic Chemistry I   
1498  2019  Spring  2019-sp    MATH     220                        Calculus   
917   2019  Spring  2019-sp    ECON     302      Inter Microeconomic Theory   
1839  2019  Spring  2019-sp    PSYC     100                     Intro Psych   
2355  2019  Summer  2019-su    PSYC     245            Industrial Org Psych   
2331  2019  Summer  2019-su     MCB     245      Human Ana
```

### Puzzle 2.3: Easiest Courses in 2019

Find the **30 easiest courses in 2019** at Illinois based on their "Average GPA".

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easiestCourses_2019 = df_2019.nlargest(30, "Average GPA")
df_easiestCourses_2019


```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
2311  2019  Summer  2019-su      IS     532    Theory & Pract Data Cleaning   
595   2019  Spring  2019-sp      CI     473           Disciplinary Literacy   
1118  2019  Spring  2019-sp    FSHN     502        Advanced Sensory Science   
1980  2019  Spring  2019-sp     SHS     352      Hearing Health and Society   
1340  2019  Spring  2019-sp     KIN     104              Skating Activities   
1911  2019  Spring  2019-sp     REL     208   Cultures & Lits of South Asia   
2362  2019  Summer  2019-su     SBC     509      Strategic Media Management   
1418  2019  Spring  2019-sp     LER     590      Gov't Reg & Employment Law   
2023  2019  Spring  2019-sp    SOCW     297       Asian Families in America   
2298  2019  Summer  2019-su    FSHN     595  Nutrition for Health & Fitness   
1310  2019  Spring  2019-sp      IS     543     Sociotechnical Info Systems   
1302  2019  Spring  2019-sp      IS     501             In
```

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_hardestCourses) == 30 ), "You should only have the 30 hardest courses."
assert( abs(df_hardestCourses["Average GPA"].mean() - 1.58) < 0.01 ), "Are you sure you have the 30 hardest courses?"

assert( len(df_hardestCourses_2019) == 30 ), "You should only have the 30 hardest courses in 2019."
assert( abs(df_hardestCourses_2019["Average GPA"].mean() - 2.04) < 0.01 ), "Are you sure you have the 30 hardest courses of 2019?"

assert( len(df_easiestCourses_2019) == 30 ), "You should only have the 30 easiest courses in 2019."
assert( abs(df_easiestCourses_2019["Average GPA"].mean() - 3.97) < 0.01 ), "Are you sure you have the 30 easiest courses of 2019?"
```

## Puzzle 3: Finding Data on  Every Subject

### Puzzle 3.1: Total Grades Given in Each Subject

Using `groupby` and `agg`, find the total number of each grade (eg: "A+"s, "A"s, "A-"s, etc.) for every subject in the DataFrame `df_grouped_subject`.

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python

df_sub = df.groupby("Subject")

df_grouped_subject = df_sub["A+", "A","A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D","D-", "F"].agg("sum").reset_index()
df_grouped_subject


```

Output:
```
Subject    A+      A     A-     B+      B    B-    C+     C   C-   D+  \
0       AAS   592   1787    755    471    446   222   121   111   45   26   
1       ABE   267   2270    748    650    851   226   121   163   46   16   
2      ACCY  6041  19837  13746  13389  12065  5064  2409  1852  918  208   
3       ACE  2466  11620   4550   3898   6123  2375  1484  2025  692  332   
4      ACES   609   1743    232     90    240    59    59    70   20   11   
..      ...   ...    ...    ...    ...    ...   ...   ...   ...  ...  ...   
156      UP   958   3203   1884   1128   1265   449   227   267  108   34   
157      VB     0    212      6      6    120     1     0    40    0    0   
158     VCM   302   1321     29     49   1036    20    18   337    1    1   
159      VM     0   1619      8      0   5283     1     0  3086    0    0   
160    YDSH    38     76     79     30     40    11     6     4    2    0   

       D   D-    F  
0     28   18   58  
1     39    5   43  
2    600  136  4
```

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_grouped_subject) == 161 ), "Are you sure you grouped by subjects?"
assert( df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["A"].sum() == 38472 ), "Are you sure you counted the right number of each subject?"
```

### Puzzle 3.2: Overall Average GPA

In `df_grouped_subject`, add a new column called `"Overall Average GPA"` that computes the total overall GPA of each subject at Illinois.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python
df_grouped_subject["Total Students"] = df_grouped_subject["A+"]+df_grouped_subject["A"]+df_grouped_subject["A-"]+df_grouped_subject["B+"]+df_grouped_subject["B"]+df_grouped_subject["B-"]+df_grouped_subject["C+"]+df_grouped_subject["C"]+df_grouped_subject["C-"]+df_grouped_subject["D+"]+df_grouped_subject["D"]+df_grouped_subject["D-"]+df_grouped_subject["F"]
df_grouped_subject["Overall Average GPA"] = (df_grouped_subject["A+"]*4.0+df_grouped_subject["A"]*4.0+df_grouped_subject["A-"]*3.67+df_grouped_subject["B+"]*3.33+df_grouped_subject["B"]*3+df_grouped_subject["B-"]*2.67+df_grouped_subject["C+"]*2.33+df_grouped_subject["C"]*2+df_grouped_subject["C-"]*1.67+df_grouped_subject["D+"]*1.33+df_grouped_subject["D"]*1+df_grouped_subject["D-"]*0.67)/ df_grouped_subject["Total Students"]
df_grouped_subject


```

Output:
```
Subject    A+      A     A-     B+      B    B-    C+     C   C-   D+  \
0       AAS   592   1787    755    471    446   222   121   111   45   26   
1       ABE   267   2270    748    650    851   226   121   163   46   16   
2      ACCY  6041  19837  13746  13389  12065  5064  2409  1852  918  208   
3       ACE  2466  11620   4550   3898   6123  2375  1484  2025  692  332   
4      ACES   609   1743    232     90    240    59    59    70   20   11   
..      ...   ...    ...    ...    ...    ...   ...   ...   ...  ...  ...   
156      UP   958   3203   1884   1128   1265   449   227   267  108   34   
157      VB     0    212      6      6    120     1     0    40    0    0   
158     VCM   302   1321     29     49   1036    20    18   337    1    1   
159      VM     0   1619      8      0   5283     1     0  3086    0    0   
160    YDSH    38     76     79     30     40    11     6     4    2    0   

       D   D-    F  Total Students  Overall Average GPA  
0     28   18   58   
```

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( abs(df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["Overall Average GPA"].sum() - 3.249) < 0.01 ), "Your overall average GPA does not appear correct."
```

### Puzzle 3.3: Find the 15 easiest subjects at Illinois

Finally, find the 15 easiest subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-cb713763b3781870, points: 0, grade: false, solution: true -->
```python
df_easiest_subjects = df_grouped_subject.nlargest(15,"Overall Average GPA")
df_easiest_subjects


```

Output:
```
Subject    A+     A    A-   B+    B   B-  C+   C  C-  D+   D  D-   F  \
34       CB     0   140     0    0   17    0   0   1   0   0   0   0   0   
155     UKR     5   190    43   20    8    4   0   1   0   0   0   0   0   
40      CHP    42   101    27    8    6    1   0   3   0   0   0   0   1   
128    REES    32    34     7    2    5    2   0   0   0   0   1   0   0   
53     EDUC   326  1788   248  112  102   25  17  29   6   2  11   2   9   
33      BUS   519  4566   944  422  232   84  31  35   9   4   9   2  23   
137     SBC   162   150    57   23   11    9   2   3   0   2   2   0   0   
89       IS  1556  6926  2062  815  566  166  43  52  15   3  19   1  34   
126     PSM   121    37    13    4   10    4   2   3   2   0   0   0   1   
11      AHS   142   708    37   21   52    4   0  14   2   0  10   0  15   
59      EOL    88   960   213  106   86   10   3   9   0   0   4   0   8   
41       CI  1429  7884  2083  944  657  175  79  86  30   7  13   1  35   
151      TE   47
```

### Puzzle 3.4: Find the hardest 15 subjects at Illinois

Find the hardest 15 subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-eb5a80d62e1c7a45, points: 0, grade: false, solution: true -->
```python
df_hardest_subjects = df_grouped_subject.nsmallest(15,"Overall Average GPA")
df_hardest_subjects


```

Output:
```
Subject     A+      A     A-     B+      B     B-     C+      C    C-  \
159      VM      0   1619      8      0   5283      1      0   3086     0   
150     TAM   1746   4130   3179   3167   3883   2576   2063   2143  1353   
105    MATH  10256  26173  16820  17217  23525  13939  11953  14394  7394   
121    PHYS  12103  14472  10042   9617  10233   9037   7444   7312  5409   
99      LAT     14     20     17     16     21      7     14      6     3   
136    SAME      2     29     30     20     31     12     14     17     5   
36     CHBE    577   3158   1593   1654   3024   1463   1188   1953   666   
38     CHIN      0      8      9      4      9      3      1      3     1   
37     CHEM  10617  40403  16307  16495  21845  12771  11475  13793  7408   
107     MCB   6114  12503   9746   9556  11556   7187   5405   5120  2954   
51      ECE   7003  20082  11291  10202  13218   7274   5024   6491  3297   
125      PS   1915   6951   5350   4438   5768   2974   1891   2070   926   
73 
```

### Puzzle 3.5: Find your major!

Finally, find your major/subject in the DataSet and check out how hard it is! 

Feel free to look around and explore! Write your code below and write comments to explain what your code does. For example, what are the hardest classes in your major that have had at least 50 students?

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python
df_major = df[df["Subject"]=="ECON"] #This creates a data frame with only ECON classes.
df_major_C = df_major.groupby("Number") #This groups together all the sections into which class they are apart of.

df_major_C_G = df_major_C["A+", "A","A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D","D-", "F"].agg("sum").reset_index()
#this will get the sum of the number of each letter grade given out for each class.

df_major_C_G["Total Students"] = df_major_C_G["A+"]+df_major_C_G["A"]+df_major_C_G["A-"]+df_major_C_G["B+"]+df_major_C_G["B"]+df_major_C_G["B-"]+df_major_C_G["C+"]+df_major_C_G["C"]+df_major_C_G["C-"]+df_major_C_G["D+"]+df_major_C_G["D"]+df_major_C_G["D-"]+df_major_C_G["F"]
df_major_C_G["Overall Average GPA"] = (df_major_C_G["A+"]*4.0+df_major_C_G["A"]*4.0+df_major_C_G["A-"]*3.67+df_major_C_G["B+"]*3.33+df_major_C_G["B"]*3+df_major_C_G["B-"]*2.67+df_major_C_G["C+"]*2.33+df_major_C_G["C"]*2+df_major_C_G["C-"]*1.67+df_major_C_G["D+"]*1.33+df_major_C_G["D"]*1+df_major_C_G["D-"]*0.67)/ df_major_C_G["Total Students"]
#First comand will get the total students who have taken each course, while the second will get the total average gpa for the course. 

df_major_C_G_50 = df_major_C_G[df_major_C_G["Total Students"]>50]
#This command will create a new data frame with classes that have had more than 50 students. 

df_major_C_G_50.nsmallest(15, "Overall Average GPA")
#This command will finally pick the 15 hardest classes that Econ has offered based on overall average GPA. 
```

Output:
```
Number    A+     A    A-    B+     B    B-    C+     C   C-   D+    D  \
4      203  1522  1481  1266  1346  1722  1233  1008  1225  609  503  479   
5      220     9     8    19    15    12     7    10     7    4    2    4   
6      302  1555  2407  1894  1697  1832  1476  1037  1054  883  512  483   
9      411    43   145   122   122   153   109    93   113   54   20   31   
1      103  1035  2835  2221  2381  2993  1613  1246  1573  694  345  560   
31     471    38   203   142   162   149   114    70    79   55   28   53   
3      202   315  2439   507   690  2580   344   479  1226  156  208  434   
21     448     1    10    15    22    18    14    11     5    3    2    0   
8      402     8    15    14     8    11     3     4     4    0    0    0   
7      303   222   568   568   679   798   488   328   296  155   82   68   
17     437    27    26    18    16     5     4     7     9    3    9    2   
12     417     8     6     6     6    14     7     5     5    3    1    0   
38 
```

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. After you have saved and checkpointed, exit this notebook, and return to the Data Science DISCOVERY page on how to use git to turn this notebook into the course!