# Lab: Hypothesis Testing


## More Sleep Always! 

SLEEP! It's a beautiful thing that most of us don't get enough of!  There has been a lot of research done looking the amount of sleep different people get on average.  The following is an example of one of them. 

In 2013, [Gallup completed their latest survey on sleep](https://news.gallup.com/poll/166553/less-recommended-amount-sleep.aspx) that shows that the average amount of sleep Americans get is 6.8 hours ($\mu = 6.8$). This seems a bit low, right? How do you think our class average in Data Science DISCOVERY compares to this? It might be less than 6.8 or it may be more than that. Let's find out if the ideal hours of sleep STAT 107 students would get is significantly different from the average amount of sleep that all Americans get reported by Gallup. For the sake of this lab, assume our survey data is a random sample from the population (Americans).

Run a **hypothesis test** to find out if the data we have is significantly different from the population/Gallup poll. 

You need to look at the Day 1 Dataset (from the first week of class) which gives you the ideal amount of sleep Stat 107 students get each week. Here's the link (make sure to login with @illinois): 
https://docs.google.com/spreadsheets/d/1uS6OoVsaCNEDwi-ozOXJe7Vn08M-41k7W1cdAlwLIiY/edit?usp=sharing

As always, start with uploading dataset to your notebook, and import the relevant libraries to your notebook.


<!-- grade_id: cell-d2f6e20542c85a37, points: 0, grade: false, solution: true -->
```python
import pandas as pd

df = pd.read_csv("Hello.csv")
df
```

Output:
```
Name                Major School Year    Phone  \
0          Wade     Computer Science       Other  Android   
1         Karle           Statistics       Other   iPhone   
2         Grace           Statistics   Sophomore   iPhone   
3       Miranda          Cs and math   Sophomore   iPhone   
4          Chen  Information Science    Freshman   iPhone   
..          ...                  ...         ...      ...   
86         Nick           Statistics   Sophomore   iPhone   
87         Howe           Statistics   Sophomore   iPhone   
88        Davin            Economics      Senior   iPhone   
89        Tanvi           Statistics   Sophomore   iPhone   
90  Christopher                  Law       Other   iPhone   

                   Computer  Holes in a Straw  Has Dogs  Has Cats  Has Fish  \
0    Windows-based computer               2.0         0         1         0   
1   Mac OS X-based computer               1.0         0         0         0   
2   Mac OS X-based computer              
```

## Puzzle 1: Sleep Among STAT 107 Students

### Part 1: Null and Alternative Hypothesis
In class, we learned that the first step when doing any hypothesis test is to state the **Null Hypothesis** and **Alternative Hypothesis**. 

**Null Hypothesis** := The null hypothesis says that any difference between our sample and the population is small and due to chance.

**Alternative Hypothesis** := The alternative hypothesis says that the differences we observe are not just due to chance. 

_Warning!!!_ Remember that when you are stating your hypotheses, they always include the **population parameters** such as $\mu$ not **sample parameters** $\bar{X}$. 

Using the Day 1 dataset and the Gallup poll, state your 2 hypotheses:

### Null Hypothesis

$H_0$ := Day1.Mean - 6.8 = 0


### Alternative Hypothesis

$H_a$ := Day1.Mean - 6.8 != 0


### Part 1a: P-Values:
In lecture, Karle and Wade discussed what a p-value is. We can use it to decide which hypothesis is supported by our data.
Here, it will tell us the probability that we will observe the difference we observed between our sample data and the Gallup data, given that the true average hours of sleep is really 6.8 hours, as reported by Gallup.

Calculate the Z-score and p-value in the following cell.  If you are using the `ztest`, the provided code will extract the z_score and p_value from `ztest_result`:

<!-- grade_id: cell-73019716433eb5e2, points: 0, grade: false, solution: true -->
```python
from statsmodels.stats.weightstats import ztest

# Run a two-sided z-test:
ztest_result = ztest(df["Hours of Sleep"],value = 6.8)



z_score, p_value = ztest_result
ztest_result
```

Output:
```
(3.545145644038473, 0.0003923962029647485)
```

<!-- grade_id: cell-aee5d44c0d1453a7, points: 1, grade: true, solution: false -->
```python
# == TEST CASES for Puzzle 1a ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert(math.isclose(z_score, 3.545145644038473)), "You z statistics does not seem correct. Make sure you pick the same formula"
assert(math.isclose(p_value, 0.0003923962029647485)), "You p-value statistics does not seem correct."
assert(z_score == ztest_result[0] and p_value == ztest_result[1]), "You did not appear to have run a ztest."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print(f"- z-score: {z_score}")
print(f"- p-value: {p_value}")
```

Output:
```
🎉 All tests passed! 🎉
- z-score: 3.545145644038473
- p-value: 0.0003923962029647485
```

### Part 1b: Confidence Intervals:
We can also use our p-value to find a range of values that likely would contain an unknown population parameter, also known as a confidence interval. Luckily, Python has a built-in function for finding confidence intervals, which takes in a data frame as its first argument, and a confidence level as the second argument.

Here is the documentation for finding a Bayesian confidence interval using Python: https://docs.scipy.org/doc/scipy-0.13.0/reference/generated/scipy.stats.bayes_mvs.html

- Using what you learned in lecture, look to understand how to use this function. :)


#### Your Program

Use the function to find a **80% confidence interval** for the average hours of sleep in the following cell.  Store the result of the `bayes_mvs` into the variable `confidence_80`:


<!-- grade_id: cell-fc9e25f78e13593c, points: 0, grade: false, solution: true -->
```python
from scipy.stats import bayes_mvs

# find a confidence interval with a level of 80% confidence
confidence_80 = bayes_mvs(df["Hours of Sleep"],0.8)


confidence_80[0]
```

Output:
```
Mean(statistic=7.217582417582417, minmax=(7.065512228419144, 7.36965260674569))
```

Now find a confidence interval with a level of 95% confidence. Notice the difference in the range of these intervals and how they change the more and more confident we become.

<!-- grade_id: cell-a2c0a263161f314d, points: 0, grade: false, solution: true -->
```python
# find a confidence interval with a level of 95% confidence
confidence_95 = bayes_mvs(df["Hours of Sleep"],0.95)


confidence_95[0]
```

Output:
```
Mean(statistic=7.217582417582417, minmax=(6.983572184744944, 7.451592650419889))
```

<!-- grade_id: cell-23cd8829704675cd, points: 1, grade: true, solution: false -->
```python
# == TEST CASES for Puzzle 1b ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math

# Min/max for `confidence_80`
assert(math.isclose(confidence_80[0][1][0], 7.065512228419144)), "Your confidence_80 does not seem correct."
assert(math.isclose(confidence_80[0][1][1], 7.36965260674569)), "Your confidence_80 does not seem correct."

# Min/max for `confidence_95`
assert(math.isclose(confidence_95[0][1][0], 6.983572184744944)), "Your confidence_95 does not seem correct."
assert(math.isclose(confidence_95[0][1][1], 7.451592650419889)), "Your confidence_95 does not seem correct."



## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print(f"- 80% interval: [{confidence_80[0][1][0]} - {confidence_80[0][1][1]}] hours of sleep")
print(f"- 95% interval: [{confidence_95[0][1][0]} - {confidence_95[0][1][1]}] hours of sleep")
print()


```

Output:
```
🎉 All tests passed! 🎉
- 80% interval: [7.065512228419144 - 7.36965260674569] hours of sleep
- 95% interval: [6.983572184744944 - 7.451592650419889] hours of sleep
```

### Part 1c: Connect the p value result with your original hypotheses :

We know what the **p-value** is now! But as a data scientist, it's important to explain your results in a way that anyone (even someone who has no background in statistics) can understand. If we just said "Here is the p-value", they may be confused. We need to be more specific about our conclusion. 

In the following cell, delete the statements that you do not believe are true and keep the one you believe is correct.


<!-- grade_id: cell-0f84432fe81139e8, points: 0, grade: false, solution: true -->
```python
# Remove 2 the of the following lines to leave only the correct answer:
#conclusion = "There is not enough evidence to reject the null hypothesis. Thus, DO NOT reject the Null"
#conclusion = "We cannot infer anything!"
conclusion = "There is enough evidence to reject the null hypothesis. Thus, DO reject the Null"



```

<!-- grade_id: cell-119c28992fed3d21, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(conclusion.split()[2][3::-1] == "uone"), "Your conclusion does not seem correct. Think about more!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 2: Who loves sleep more?

We are curious if there is a significant difference in ideal hours of sleep between introverts and ambiverts in college. We are specifically interested in whether introverts sleep **less than** extroverts.

We can investiage this using our Day 1 Dataset as well!  Here, assume that the Day 1 Dataset is a random sample taken from all college students.

Do a hypothesis test to satisfy our curiosity! As always, first you need to state your hypotheses.

### Part 2a: Null and Alternative Hypothesis

### Null Hypothesis

$H_0$ := Introvert.Mean = Extrovert.Mean

### Alternative Hypothesis

$H_a$ := Introvert.Mean < Extrovert.Mean

### Part 2b: P-value:

Calculate your test statistic and find the p-value. Remember, we are comparing two samples now- think about how this is different from Puzzle 1 and what type of z-test to do!

Find the **p-value** in the following cell: 

<!-- grade_id: cell-6c8eb2d2b6927e26, points: 0, grade: false, solution: true -->
```python


df_I = df[df['Personality']=="Introvert"]
df_E = df[df['Personality']=="Extrovert"]



# Run a two-sample ztest comparing introverts and extrverts:
ztest_result = ztest(df_I["Hours of Sleep"],df_E["Hours of Sleep"],value = 0)



z_score = ztest_result[0]
p_value = ztest_result[1]
ztest_result
```

Output:
```
(-0.7208202957997942, 0.47102008654087113)
```

<!-- grade_id: cell-7bd5b9e9b0bf72b8, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert(math.isclose(z_score, -0.7208202957997942)), "You z statistics does not seem correct."
assert(math.isclose(p_value, 0.47102008654087113)), "You p-value statistics does not seem correct."
assert(z_score == ztest_result[0] and p_value == ztest_result[1]), "You did not appear to have run a ztest."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print(f"- z-score: {z_score}")
print(f"- p-value: {p_value}")
```

Output:
```
🎉 All tests passed! 🎉
- z-score: -0.7208202957997942
- p-value: 0.47102008654087113
```

### Part 2c: Make a conclusion based on your p-value.

We know what the **p-value** is now. So in the following cell delete the statements that you do not believe are true and keep the one you think is true.  Use 1% as your cutoff instead of 5%.

<!-- grade_id: cell-5f2236ee09b8ae71, points: 0, grade: false, solution: true -->
```python
# Remove the of the following lines to leave only the correct answer:
#conclusion = "There is enough evidence to reject the null hypothesis. Thus, DO reject the Null"
#conclusion = "We cannot infer anything!"
conclusion = "There is not enough evidence to reject the null hypothesis. Thus, DO NOT reject the Null"


```

<!-- grade_id: cell-5ba76e5cb1428aa4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(conclusion.split()[2][3::-1]  == "ton"), "Your conclusion does not seem correct. Think about more!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

Output:
```
🎉 All tests passed! 🎉
```

### Puzzle 3: GPA Visualization

Using Wade's [Grade disparity between sections at UIUC](http://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/) visualization, find the average grade Wade gave to CS 225 students.

You want to check if Wade was completely honest with his grades, so you randomly survey 107 students who have taken CS 225.  The final grades the students gave you is out of 1000 points.  You have recorded the 107 student grades in `cs225-final-grades.csv`.  (*These are really 107 randomly chosen grades from CS 225 from Fall 2018.*)

You find that the translation between points and grades is as follows:

| Points | Grade | GPA |
| ------ | ----- | --- |
| 1070 - 1100 | A+ | 4.0 |
|  930 - 1070 | A  | 4.0 |
|  900 - 930 | A-  | 3.67 |
|  870 - 900 | B+  | 3.33 |
|  830 - 870 | B   | 3.0 |
|  800 - 830 | B-  | 2.67 |
|  770 - 800 | C+  | 2.33 |
|  730 - 770 | C   | 2.0 |
|  700 - 730 | C-  | 1.67 |
|  670 - 700 | D+  | 1.33 |
|  630 - 670 | D   | 1.0 |
|  600 - 630 | D-  | 0.67 |
|    0 - 600 | F  | 0 |

**Hint**: You will need to process `cs225-final-grades.csv` and understand and gather data from the GPA visualization before you will be able to run the test.  Use `pandas` to help you process this data! :)

```python

```

<!-- grade_id: cell-053f4c444b109148, points: 0, grade: false, solution: true -->
```python
# First read the data cs225-final-grades.csv
sample_cs225 = pd.read_csv("cs225-final-grades.csv")

# Procecess sample_cs225 to convert to GPA format. 
# - The `sample_cs225` must contain a GPA column with the GPA as a numeric value.





for i in range(107):
    x = sample_cs225.loc[i, 'CS 225 Final Grade']
    
    if x >= 930:
        sample_cs225.loc[i, 'GPA'] = 4
    elif x >= 900 and x< 930:
        sample_cs225.loc[i, 'GPA']= 3.67
    elif x >=870 and x< 900:
        sample_cs225.loc[i, 'GPA']= 3.33
    elif x >= 830 and x< 870:
        sample_cs225.loc[i, 'GPA']= 3.00
    elif x >=800 and x< 830:
        sample_cs225.loc[i, 'GPA']= 2.67
    elif x >=770 and x< 800:
        sample_cs225.loc[i, 'GPA']= 2.33
    elif x >=730 and x< 770:
        sample_cs225.loc[i, 'GPA']= 2.00
    elif x >=700 and x< 730:
        sample_cs225.loc[i, 'GPA']= 1.67
    elif x >=670 and x< 700:
        sample_cs225.loc[i, 'GPA']= 1.33
    elif x >=630 and x< 670:
        sample_cs225.loc[i, 'GPA']= 1.00
    elif x >=600 and x< 630:
        sample_cs225.loc[i, 'GPA']= 0.67
    else:
        sample_cs225.loc[i, 'GPA']= 0
sample_cs225.head(5)
len(sample_cs225)
```

Output:
```
107
```

<!-- grade_id: cell-0544cc08b9b64325, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3a ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(sample_cs225[sample_cs225["GPA"] == 3.67]) == 7), "Your convertion from the total points to GPA does not seem correct."
assert(len(sample_cs225[sample_cs225["GPA"] == 0]) == 7), "Your convertion from the total points to GPA does not seem correct."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

### Part 3a: Hypotheses

Just like before, start with stating your Null and Alternative Hypotheses!

### Null Hypothesis

$H_0$ := sample.AVG - pop.AVG = 0

### Alternative Hypothesis

$H_a$ := sample.AVG - pop.AVG != 0

### Part 3b: P-value
Calculate the test statistic and find the p-value in the following cell:

<!-- grade_id: cell-26dab0ff60b63979, points: 0, grade: false, solution: true -->
```python


# Run a one-sample z-test
ztest_result = ztest(sample_cs225['GPA'],value = 3.22)



z_score = ztest_result[0]
p_value = ztest_result[1]
ztest_result
```

Output:
```
(1.5016743453565295, 0.1331812325552574)
```

<!-- grade_id: cell-08a0447915dbc44c, points: 1, grade: true, solution: false -->
```python
# == TEST CASES for Puzzle 3b ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(round(z_score,1) == 1.5), "You z statistics does not seem correct. Make sure you pick the same formula"
assert(round(p_value,1) == 0.1), "You p-value statistics does not seem correct."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

## Submit Your Work!

You're almost done -- congratulations!

You need to do a few more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!